# Homework 3: Files, Objects, and Functional Programming (38 points)

```python
name: Ying Jiang

email: jiangyyy@umich.edu

I did not discuss this homework with anyone.

This homework assignment took me 12 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)
```

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. Before submitting, please make sure make sure __Cells->Run__ executes without errors; errors in your code cause problems for our auto-grader. Thus, even though we may ask you to check for errors in your functions, you should not include in your submission any examples of your functions actually raising those errors.

<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not comprehensive, additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all relevant cases. E.g. If the input to a function are positive integers, then it should work for all positive integers."

**Note:** Error checking is not required unless explicitly stated in the problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:

```python
%load_ext pycodestyle_magic
%flake8_on --ignore E266
```


## Problem 1: Counting Word Bigrams (6 points) 
In your previous homework, you wrote a function for counting character bigrams. Now, let’s write a function for counting word bigrams. That is, for each pair of words, say, cat and dog, we want to count how many times the word “cat” occurred immediately before the word “dog”. We will represent this bigram by a tuple, `(’cat’, ’dog’)`. 

So, as an example, the fragment of poem,

```
Half a league, half a league, Half a league onward,
All in the valley of Death Rode the six hundred.
```

includes the bigrams `(’half’, ’a’)` and `(’a’, ’league’)` both three times, the bigram `(’league’, ’half’)` appears twice, while the bigram `(’in’, ’the’)` appears only once.

**Note:** For our purposes, we will ignore all spaces, newlines, punctuation and capitalization in our counting.

### Part 1 (2 points) 
Write a function `count_bigrams_in_file` that takes a ﬁlename `file` as its only argument. Your function should read from the given ﬁle, and return a dictionary whose keys are bigrams (given in the tuple form above), and values are the counts for those bigrams. 

Again, your function should ignore punctuation, spaces, newlines and capitalization.

**Hint:** you will ﬁnd the Python function `str.strip()`, along with the string constants deﬁned in the string documentation (https://docs.python.org/3/library/string.html), useful in removing punctuation. You are free to use the `re` library for this problem but our solution does not use it. 

**Hint:** be careful to check that your function handles newlines correctly. For example, in the poem above, one of the `(’league’, ’half’)` bigrams spans a newline, but should be counted nonetheless. 

**Hint:** For this function to be executed at a reasonable speed, you should attempt to remove all punctuation from the file in the text __before__ splitting it into words. For reference, our solution only uses one for loop over all the words in the given file's text. 

**Note:** be careful that your function does not accidentally count the empty string as a word (this is a common bug if you aren’t careful about splitting the input text). Solutions that merely delete “bad” keys from the dictionary at the end will not receive full credit, as all edge cases can handled by correctly splitting the input.

**Note:** Please replace instances of '--' with a space ' '. So that e.g. 'Bob--David' would be parsed as 'Bob David' (two separate words). Any other punctuation should be removed and treated the same, so for example 'Bob/David' would be parsed as 'BobDavid'. This includes punctuation inside words so e.g. "they're" would become "theyre"

<!-- grade_id: cell-00acb788c373a75a, points: 0, grade: false, solution: true -->
```python
import string

def count_bigrams_in_file(file):
    bigrams=[('a', 'league'),('half', 'a'),('league', 'half'),('league', 'onward')]#list of bigrams
    content=''
    count=dict()
    
    #get the words
    file=open(file)
    content=file.read() 
    content=content.lower()
    content=content.replace('--'," ")
    for i in content:
        if i in string.punctuation:  
            content = content.replace(i,"")# the replace function
    content=content.split()
    
    for i in range(len(content)-1):
        if (content[i],content[i+1]) in bigrams:
            if (content[i],content[i+1]) not in count.keys():
                count[(content[i],content[i+1])]=1
            else:
                count[(content[i],content[i+1])]+=1
            
    return(count)
```

<!-- grade_id: cell-64a1f6983f806530, points: 2, grade: true, solution: false -->
```python
"""Test simple cases"""
sol = {('a', 'league'): 3, ('half', 'a'): 3,
       ('league', 'half'): 2, ('league', 'onward'): 1}
with open('./test.txt', 'w') as f:
    f.write('Half a league, half a league,\n Half a league onward.')
assert count_bigrams_in_file('test.txt') == sol
```

### Part 2 (1 point)  
The ﬁle `WandP.txt`, which was included along with this jupyter notebook in `hw3.zip`, is an `ASCII` copy of all of Tolstoi’s novel War and Peace. Run your function on this ﬁle, and pickle the resulting dictionary in a ﬁle called `mb.bigrams.pickle`.

<!-- grade_id: cell-957e399ee02c11d9, points: 0, grade: false, solution: true -->
```python
"""Code for creating 'mb.bigrams.pickle' here."""
import pickle
filename='WandP.txt'
file=open("mb.bigrams.pickle","wb")
file.write(pickle.dumps(count_bigrams_in_file(filename)))
file.close()
```

<!-- grade_id: cell-43238c586b58202d, points: 1, grade: true, solution: false -->
```python
"""Check that your pickled file is exists and is correct."""
import os
import pickle

assert os.path.isfile('mb.bigrams.pickle')
```

### Part 3 (2 points)

We say that word `A` is _collocated_ with word `B` in a text if words `A` and `B` occur immediately one after another (in either order). That is, words `A` and `B` are collocated if and only if either of the tuples `(A, B)` or `(B, A)` are present in the text. 

Write a function `collocations` that takes a ﬁlename `file` as its only argument and returns a dictionary. Your function should read from the given ﬁle and return a dictionary whose keys are all the strings appearing in the ﬁle and the value of word `A` is a Python set containing all the words collocated with `A`. 

Follow the same conventions as Part 1/2 in ignoring case and stripping away all spaces, newlines and punctuation.

Again using the poem fragment above as an example, the string `’league’` should appear as a key, and should have as its value the set `{’a’, ’half’, ’onward’}`, while the string `’in’` should have the set `{’all’, ’the’}` as its value. 

**Hint:** we didn’t discuss Python sets in lecture, because they are essentially just dictionaries without values. See the documentation at https://docs.python.org/3/tutorial/datastructures.html#sets for more information.

**Hint:** For this function to be executed at a reasonable speed, you should attempt to remove all punctuation from the file in the text __before__ splitting it into words. For reference, our solution only uses one for loop over all the words in the given file's text. 

<!-- grade_id: cell-9a8dbd9611c8eb06, points: 0, grade: false, solution: true -->
```python
def collocations(file):
    file=open(file)
    #get the words
    content=file.read() 
    content=content.lower()
    content=content.replace('--'," ")
    for i in content:
        if i in string.punctuation:  
            content = content.replace(i,"")
    content=content.split()
   
   
    codict=dict()
    #find the next
    for i in range(0,(len(content)-1)):
        if content[i] not in codict.keys():
            codict[content[i]]=[content[i+1]]
        else:
            if content[i+1] not in codict[content[i]]:
                codict[content[i]].append(content[i+1])
    #find the former
    for i in range(1,len(content)):
        if content[i] not in codict.keys():
            codict[content[i]]=[content[i-1]]
        if content[i-1] not in codict[content[i]]:
            codict[content[i]].append(content[i-1])
    codict1=dict()
    for key,value in codict.items():
        codict1[key]=set(value)

    #special condition with len(content)=1
    if len(content)==1:
        codict1[content[0]]={}
    return codict1
```

<!-- grade_id: cell-9fa29d301a3567b5, points: 2, grade: true, solution: false -->
```python
"""Test it works on simple case."""
with open('test.txt', 'w') as f:
    f.write('Half a league, half a league,\n Half a league onward.')
assert collocations('test.txt') == {
                             'half': {'league', 'a'},
                             'a': {'half', 'league'},
                             'league': {'onward', 'half', 'a'},
                             'onward': {'league'}
                              }
```

### Part 4 (1 point) 
Run your function on the ﬁle `WandP.txt` and pickle the resulting dictionary in a ﬁle called `mb.colloc.pickle.`

<!-- grade_id: cell-49a3dcf4ea8de03f, points: 0, grade: false, solution: true -->
```python
"""Code for creating 'mb.colloc.pickle'."""
import pickle
filename='WandP.txt'
file=open('mb.colloc.pickle',"wb")
file.write(pickle.dumps(collocations(filename)))
file.close()
```

<!-- grade_id: cell-04c1ab0c9d863a66, points: 1, grade: true, solution: false -->
```python
"""Check that your pickled file exists and is correct"""
import os
import pickle
assert os.path.exists('mb.colloc.pickle')
```

## Problem 2: More Fun with Vectors (12 points)

In this exercise, we’ll encounter our old friend the vector yet again, this time taking an object-oriented approach.

**Note:** Before starting this problem, you are strongly encouraged to read _Chapter 14: Python Objects_ in the textbook: https://www.py4e.com/html3/14-objects

<!-- grade_id: cell-5b18f2366afe4b38, points: 0, grade: false, solution: true -->
```python
"""Include all the code for the Vector class here."""
import math


class Vector:
    dimension=0
    entries=None
    def __init__(self, dim=0, vec=None):
        if not isinstance(dim, int):
            raise TypeError
        elif dim<0:
            raise ValueError
        # dim=0 
        if dim == 0:
            if vec != None and len(vec)!=0:
                raise ValueError
            else:
                self.dimension = dim
                self.entries = tuple()
        else:
            if not isinstance(vec, (list, tuple)):
                raise TypeError
            if not all(isinstance(i, (int, float)) for i in vec):
                raise ValueError
            if len(vec)!=dim:
                raise ValueError
            self.dimension = dim
            self.entries = tuple(vec)
        
    #equality, less than, less or equal to, greater than,
    def __ne__(self, other):
        return(self.entries != other.entries)
    def __eq__(self,other):
        return( self.entries==other.entries )
    def __lt__(self,other):
        return(self.entries<other.entries)
    def __gt__(self,other):
        return(self.entries>other.entries)
    def __ge__(self, other):#大于等于
        return(self.entries>=other.entries)
    def __le__(self, other):#小于等于
        return(self.entries<=other.entries)
    #dot method
    def dot(self,other):
        sum=0
        for i in range(len(self.entries)):
            sum =sum+ self.entries[i]*other.entries[i]
        return(sum)
    
    #multiple method
    def __mul__(self,other):
        product=list()
        if isinstance(other,(int,float)):
            for value in self.entries:
                product.append(other*value)
        elif isinstance(other,Vector):
            if self.dimension!=other.dimension: 
                raise ValueError
            else:
                for i in range(self.dimension):
                    product.append(self.entries[i]*other.entries[i])
        else: return TypeError
        return(Vector(self.dimension,product))
    
    def __rmul__(self,other):
        return self.__mul__(other)
    
    def norm(self,p):
        value=0
        if p==0:value=sum(int(self.entries[i]!=0) for i in range(self.dimension))
        if p>0 and p<math.inf:
            value=(sum(abs(self.entries[i])**p for i in range(self.dimension)))**(1/p)
        if p==math.inf:
            absvalue=[abs(self.entries[i]) for i in range(self.dimension)]
            value=max(absvalue)
        return(value)
```

### Part 1 (5 points)
Deﬁne a class `Vector`. The initializer for your class should take the dimension as its ﬁrst argument and a `list` or `tuple` of numbers (`int`s or `ﬂoat`s), representing the vector’s entries, as its second argument (**Note:** your initializer should work correctly given either a list or a tuple of numbers). The initializer should create a vector with dimension $0$ and no entries (an empty tuple `tuple()`) if the user does not provide the corresponding arguments.

Your class vector should have two __instance__ attributes: `dimension` and `entries`. The first containing the dimension of the vector as an `int` and the second containing the entries of the vector as a `tuple`. 

The initializer should raise an appropriate error in the case where the dimension is invalid (i.e., wrong type or a negative number), and should also raise an error in the event that the dimension and the number of supplied entries disagree.

**Note:** Please raise `ValueError` or `TypeError` for your error checking in this problem. Raising either will work. 

<!-- grade_id: cell-d84d85fc86863cd5, points: 1, grade: true, solution: false -->
```python
"""Check initialization works on sensible cases"""
v = Vector(4, (1, 2, 3, 4))
v = Vector(4, [1, 2, 3, 4])
v = Vector(1, [1])
v = Vector(2, [2, 3])
v = Vector(1, (1,))
v = Vector(0, [])
v = Vector(0)
assert True
```

<!-- grade_id: cell-ee5b2851ac09f0d4, points: 1, grade: true, solution: false -->
```python
"""Vector has correct attributes"""
v = Vector(4, (1, 2, 3, 4))
assert isinstance(v, Vector)
assert 'entries' in v.__dict__.keys()
assert 'dimension' in v.__dict__.keys()
assert type(v.dimension) == int
assert type(v.entries) == tuple
```

<!-- grade_id: cell-556867b4cb67ffaf, points: 1, grade: true, solution: false -->
```python
"""Check you can get entries from vector and they are correct when retrieved"""
v = Vector(3, [1, 2, 3])
assert tuple(v.entries) == (1, 2, 3)
```

<!-- grade_id: cell-9eac4809b5a757a5, points: 2, grade: true, solution: false -->
```python
"""Check exceptions are handled correctly"""
try:
    Vector(5, (1, 2, 3, 4))
    assert False
except ValueError:
    assert True

try:
    Vector(-4, (1, 2, 3, 4))
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector('a', [1, 2, 3, 4])
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector(1, [])
    assert False
except (ValueError, TypeError):
    assert True
```

### Part 2 (2 point)
Implement the necessary operator(s) to support comparison (equality, less than, less or equal to, greater than, etc) of `Vector` objects. We will say that two `Vector` objects are equivalent if they have the same coordinates. Otherwise, comparison should be analogous to `tuple`s in Python, so that comparison is done on the ﬁrst coordinate ﬁrst, then the second coordinate, then the third, and so on. So, for example, the two-dimensional vector `(2, 4)` is ordered before (less than) `(2, 5)`. 

You can assume you are always comparing vectors of the same dimension. 

<!-- grade_id: cell-3a1c2f17cf2a4779, points: 1, grade: true, solution: false -->
```python
assert Vector(3, [2, 1, 1]) != Vector(3, [1, 3, 3])
assert Vector(3, [2, 1, 1]) >= Vector(3, [1, 3, 3])
assert Vector(3, [1, 1, 1]) == Vector(3, [1, 1, 1])
```

<!-- grade_id: cell-de890d7da24aa5f0, points: 1, grade: true, solution: false -->
```python
assert Vector(3, [1, 1, 1]) != Vector(3, [1, 1, 3])
assert Vector(3, [1, 3, 3]) <= Vector(3, [2, 1, 1])
```

### Part 3 (1 point) 
Implement a method `Vector.dot` that takes a single Vector as its argument and returns the inner product of the caller with the given `Vector` object. You can assume the input is always a vector of the same dimension as the original one.

<!-- grade_id: cell-2e45b3413aad5e1c, points: 1, grade: true, solution: false -->
```python
"""Check your Vector.dot() method works correctly"""
assert Vector(3, (1, 2, 3)).dot(Vector(3, (1, 2, 2))) == 11.0
```

### Part 4 (2 points)
We would also like our Vector class to support scalar multiplication. Left- or right multiplication by a scalar, e.g., `2*v` or `v*2`, where `v` is a Vector object, should result in a new `Vector` object with its entries all scaled by the given scalar. We will also follow `R` and `numpy` (which you will learn in a few weeks), and use `*` to denote entrywise vector-vector multiplication, so that for `Vector` objects `v` and `w`, `v*w` results in a new `Vector` object, with the i-th entry of `v*w` equal to the i-th entry of `v` multiplied by the i-th entry of `w`. Implement the appropriate operators to support this multiplication operation. 

Many languages have a convention for dealing with multiplication of vectors that diﬀer in their dimension, but we will not handle this case. Your method should raise a `ValueError` in the event that `v` and `w` disagree in their dimensions.

<!-- grade_id: cell-a5f2178bd7c9f252, points: 1, grade: true, solution: false -->
```python
assert 2 * Vector(3, (1, 2, 3)) == Vector(3, (2, 4, 6))
assert Vector(3, (1, 2, 3)) * 2 == Vector(3, (2, 4, 6))
assert Vector(3, (1, 2, 3)) * Vector(3, (1, 2, 3)) == Vector(3, (1, 4, 9))
```

<!-- grade_id: cell-cd6422e310c07ed2, points: 1, grade: true, solution: false -->
```python
try:
    Vector(3, (1, 2, 3)) * Vector(2, (1, 2))
    assert False
except ValueError:
    assert True
```

### Part 5 (2 points)
For a real number $0 \leq p \leq \infty$, and a vector $v \in \mathbb{R}^{d}$, the $p$-norm of $v$, written $\| v \|_{p}$, is given by: 

$$ 
\| v \|_p = \begin{cases}
                        \sum_{i=1}^d 1_{v_i \neq 0} &\mbox{ if } p = 0\\
                        (\sum_{i=1}^d |v_i|^p )^{1/p} &\mbox{ if }
                                                        0 < p < \infty, \\
                        \max_{i=1,2,\dots,d} |v_i| &\mbox{ if } p=\infty
                        \end{cases}
$$

where $v_i$ is the $i$-th component of the vector and the notation $1_{v_i \neq 0}$ indicates: 

$$ 
1_{v_i \neq 0} = \begin{cases}
                        1 &\mbox{ if } v_i \neq 0\\
                        0 &\mbox{ if } v_i = 0
                        \end{cases}. 
$$


Implement a method `Vector.norm` that takes a single int or ﬂoat `p` as an argument and returns the $p$-norm of the calling `Vector` object. You can assume the input $p$ is always a non-negative integer or float.

`Hint:` see https://docs.python.org/3/library/functions.html# float for documentation on representing positive inﬁnity in Python.

<!-- grade_id: cell-2af28a606f1af266, points: 1, grade: true, solution: false -->
```python
assert Vector(3, (1, 2, 3)).norm(1) == 6.0
assert Vector(2, (1, 2)).norm(2) == math.sqrt(1**2 + 2**2)
```

<!-- grade_id: cell-a8f312c44b5334d0, points: 1, grade: true, solution: false -->
```python
assert Vector(5, (3, 2, 1, 10, 5)).norm(math.inf) == 10
assert Vector(5, (3, 2, 1, 10, 5)).norm(0) == 5
```

##  Problem 3: Iterators and Generators (7 points) 
In this exercise, you’ll get some practice working with iterators and generators. 

**Note:** in this problem, the word enumerate is meant in the sense of returning elements, not in
the sense of the Python function enumerate. So, if I say that "an iterator enumerates a
sequence $a_{0}, a_{1}, a_{2}, ...$", I mean that these are the elements that it returns upon calls to
the `__next__` method, not that it returns pairs $(i,a_{i})$ like the `enumerate` function.

### Part 1 (1 point) 
Deﬁne a class `Fibo` of iterators that enumerate the Fibonacci numbers. For the purposes of this problem, the Fibonacci sequence begins $0, 1, 1, 2, 3, ...$ with the $n$-th Fibonacci number $F_{n}$ given by the recursive formula $F_{n} = F_{n-1} + F_{n-2}$  where $F_{0} = 0$ and $F_{1} = 1$. 

Your solution should not make use of any function aside from addition (i.e., you should not need to use the function `fibo()` deﬁned in lecture a few weeks ago). Your class should support, at a minimum, an initialization method, a `__iter__` method (so that we can get an iterator) and a `__next__` method. 

**Note:** there is an especially simple solution to this problem that can be expressed in just a few lines using tuple assignment.

<!-- grade_id: cell-bbaefcc15e4a7081, points: 0, grade: false, solution: true -->
```python
class Fibo():
    def __init__(self):
        self.f0=0
        self.f1=1
        self.pos=0
    def __iter__(self):
        return self
    
    def __next__(self):
        fibo=self.f0
        self.f0,self.f1=self.f1,self.f0+self.f1
        self.pos +=1
        return(fibo)
```

```python
class Fibo():
    def __init__(self):
        self.f0=0
        self.f1=1
    def __iter__(self):
        return self
    def __next__(self):
        result = self.f0
        (self.f0, self.f1)=(self.f1, self.f1+self.f0)
        return (result)
```

<!-- grade_id: cell-e98351d056391b0a, points: 1, grade: true, solution: false -->
```python
ans1 = Fibo()
assert [next(ans1) for _ in range(10)] == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34]
```

### Part 2 (2 points)
We can generalize the Fibonacci sequence by following the same recursive procedure $F_{n} = F_{n-1}+F_{n-2}$ , but using a diﬀerent choice of initial two values for $F_{0}$ and $F_{1}$. For example, if we take $F_{0} = 2$ and $F_{1} = 1$, then we obtain the Lucas numbers, which are closely related to the Fibonacci numbers (https://en.wikipedia.org/wiki/Lucas_number). 

Deﬁne a class `GenFibo` of iterators that enumerate generalized Fibonacci numbers. Your class should inherit from the Fibo class deﬁned in the previous subproblem. The initialization method for the GenFibo class should take two optional arguments that specify the values of $F_{0}$ and $F_{1}$ , in that order, and their values should default so that `F = GenFibo()` results in an enumerator equivalent to the one that would have been created if you had called `F = Fibo()` (i.e., `GenFibo()` should produce an iterator over the Fibonacci numbers).

```python
class GenFibo(Fibo):
    def __init__(self, a0=0, a1=1):
        self.f0=a0
        self.f1=a1
```

<!-- grade_id: cell-2cea9a637e433f6a, points: 0, grade: false, solution: true -->
```python
class GenFibo(Fibo):
    def __init__(self, a0=0, a1=1):
        self.f0=a0
        self.f1=a1
        self.pos=0
```

<!-- grade_id: cell-62d3630b587beb0d, points: 1, grade: true, solution: false -->
```python
ans = GenFibo(2, 1)
assert [next(ans) for _ in range(10)] == [2, 1, 3, 4, 7, 11, 18, 29, 47, 76]
```

<!-- grade_id: cell-3b1263325c0e043d, points: 1, grade: true, solution: false -->
```python
ans = GenFibo()
assert [next(ans) for _ in range(10)] == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34]
```

### Part 3 (2 point) 

Deﬁne a generator `primes` that enumerates the prime numbers. Recall that a prime number is any integer $p > 1$ whose only divisors are $p$ and $1$. 

**Note:** you may use the function `is_prime` that we deﬁned in class (or something similar to it) and receive full credit. But there is a more graceful solution that avoids declaring a separate function or method for directly checking primality. 

**Hint:** consider a pattern similar to the one seen in lecture using the any and/or all functions.

```python
import math
def primes():
    n=2
    yield n
    while 1:
        n += 1
        if all(n%i !=0 for i in range(2,n)):
            yield n
```

<!-- grade_id: cell-6075e6fd4a6029d8, points: 0, grade: false, solution: true -->
```python
import math

def primes():
    n=2
    yield n
    while 1:
        n=n+1
        if not any(n%x==0 for x in range(2,n)):
            yield n;
```

<!-- grade_id: cell-39a0a188d4945486, points: 1, grade: true, solution: false -->
```python
ans = primes()
assert [next(ans) for _ in range(10)] == [2, 3, 5, 7, 11, 13, 17, 19, 23, 29]
```

<!-- grade_id: cell-b51f7843f8591a1f, points: 1, grade: true, solution: false -->
```python
ans = primes()
assert [next(ans) for _ in range(1)] == [2]
ans = primes()
assert [next(ans) for _ in range(0)] == []
```

### Part 4 (2 points) 
This one is good practice for coding interview questions. The _Ulam numbers_ are a sequence $u_1 , u_2 , u_3, ...$ of positive integers, deﬁned in the following way: $u_1 = 1$, and $u_2 = 2$. For all $n > 2$, $u_n$ is the smallest integer that is expressible as a sum of two distinct terms from earlier in the sequence in exactly one way. See the Examples section of the Wikipedia page for an illustration: https://en.wikipedia.org/wiki/Ulam_number. 

Deﬁne a generator `ulam` that enumerates the Ulam numbers. 

**Hint:** it will be helpful to try and break this problem into smaller, simpler subproblems. In particular, you may ﬁnd it helpful to write a function that takes a list of integers `t` and one additional integer `u`, and determines whether or not `u` is expressible as a sum of two distinct elements of `t` in exactly one way.

<!-- grade_id: cell-ae05ce7df31047af, points: 0, grade: false, solution: true -->
```python
import itertools
def is_ulam(t, n):#in exactly one way
    allcombine=list(itertools.combinations(t,2))
    alldict=dict()
    for (i,j) in allcombine:
        if alldict.get(i+j)==None: alldict[i+j]=[(i,j)]
        else: alldict[i+j].append((i,j))
    alln=[key for key in alldict.keys() if len(alldict[key])==1]
    
    return (n in alln)


def ulam():
    u1=1;u2=2
    yield u1
    yield u2
    useq=[u1,u2]
    n=3
    while 1:
        if is_ulam(useq, n):
            useq.append(n)
            yield n
        n=n+1
```

<!-- grade_id: cell-88e21fdf334740cf, points: 1, grade: true, solution: false -->
```python
ans = ulam()
assert [next(ans) for _ in range(18)] == [1, 2, 3, 4, 6, 8, 11, 13, 16, 18,
                                          26, 28, 36, 38, 47, 48, 53, 57]
```

<!-- grade_id: cell-f869c2e123a88ffc, points: 1, grade: true, solution: false -->
```python
ans = ulam()
assert [next(ans) for _ in range(5)] == [1, 2, 3, 4, 6]
```

##  Problem 4: List Comprehensions and Generator Expressions (5 points)
In this problem you’ll write a few simple list comprehensions and generator expressions. Again in this problem I use the term _enumerate_ to mean that a list comprehension or generator expression returns certain elements, rather than in the sense of the Python function `enumerate`.

### Part 1 (1 point) 
Write a list comprehension that enumerates the sequence $3^{n}−1$ for $n = 1, 2, 3,..., 20$.

For ease of grading, please assign this list comprehension to a variable called `pow3minus1`.

<!-- grade_id: cell-5453d98b9528fcce, points: 0, grade: false, solution: true -->
```python
"""Write the list comprehension here. Make sure to assign it
to a variable `pow3minus1`"""
# YOUR CODE HERE
pow3minus1=list(map(lambda x:3**x-1,range(1,21)))
```

<!-- grade_id: cell-3bb76ac1abf8f11e, points: 1, grade: true, solution: false -->
```python
"""Check your pow3minus1 is correct."""
assert pow3minus1[0:2] == [2, 8]
```

### Part 2 (2 points) 
The _Lazy Caterer’s sequence_ is a sequence of numbers that counts, for each $n = 0, 1, 2, ...$ the largest number of pieces that can be cut from a disk with at most $n$ cuts (https://en.wikipedia.org/wiki/Lazy_caterer’s_sequence). The $n$-th number in this sequence is given by $p_{n} = (n^{2} + n + 2)/2$, where $n = 0, 1, 2, ...$

Write a generator expression that enumerates the Lazy Caterer’s sequence. For ease of grading, please assign this generator expression to a variable called `caterer`. 

**Hint:** you may ﬁnd it useful to deﬁne a generator that enumerates the non-negative integers.

<!-- grade_id: cell-305df1e58b315577, points: 0, grade: false, solution: true -->
```python
"""Your code goes here for creating the assigning `caterer` generator."""
"""Your code goes here for creating the assigning `caterer` generator."""
def nonnegint():
    n= -1
    while True:
        n += 1
        yield n
n = nonnegint()

caterer = ((i**2 + i + 2)/ 2 for i in n)
```

<!-- grade_id: cell-c450f6f17957115d, points: 1, grade: true, solution: false -->
```python
"""Check your caterer generator is correct"""
caterer_list = [next(caterer) for _ in range(10)]
assert caterer_list[0:2] == [1, 2]
```

<!-- grade_id: cell-b8ef589f41a34ec3, points: 1, grade: true, solution: false -->
```python
"""Check your caterer generator is correct"""
assert caterer_list[0:5] == [1, 2, 4, 7, 11]
```

### Part 3 (1 point) 
Write a generator expression that enumerates the pyramid numbers. The $n$-th pyramid number $P_{n}$ (for $n = 1, 2, ...$) counts the number of spheres in a pyramid with an $n$-by-$n$ based (see https://en.wikipedia.org/wiki/Square_pyramidal_number), and is given by:
$$
P_{n} = \sum_{k=1}^{n}k^2 = \frac{n (n+1) (2n+1)}{6}
$$
For ease of grading, please assign this generator expression to a variable called `pyramid`. 

**Hint:** you may ﬁnd it useful to deﬁne a generator that enumerates the positive integers.

<!-- grade_id: cell-191508c2cdde347d, points: 0, grade: false, solution: true -->
```python
"""Your code for the pyramid generator goes here."""
def posint():
    n=0
    while 1:
        n +=1
        yield n
n=posint()
pyramid=(i*(i+1)*(2*i+1)/6 for i in n)
```

Output:
```
2:1: E302 expected 2 blank lines, found 0
3:6: E225 missing whitespace around operator
5:13: E225 missing whitespace around operator
7:1: E305 expected 2 blank lines after class or function definition, found 0
7:2: E225 missing whitespace around operator
8:8: E225 missing whitespace around operator
```

<!-- grade_id: cell-8eb2d2188e8b54d2, points: 1, grade: true, solution: false -->
```python
"""Check your pyramid generator is correct"""
pyramid_list = [next(pyramid) for _ in range(10)]
assert pyramid_list[:5] == [1, 5, 14, 30, 55]
```

### Part 4 (1 point) 
Write a generator expression that enumerates the _octahedral numbers_. The $n$-th octahedral number ($n = 1, 2, ...$) is given by 

$$O_{n} = \frac{n(2n^{2}+1)}{3},$$ 

and counts the number of spheres in an octahedron with $n$ spheres to each edge (see https://en.wikipedia. org/wiki/Octahedral_number). 

For ease of grading, please assign this generator expression to a variable called `octa`. 

**Note:** a particularly clever solution to this problem could take advantage of the fact that the $n$-th octahedral number can be expressed as $O_{n} = P_{n} + P_{n-1}$ where $P_{n}$ denotes the $n$-th pyramidal number, which you dealt with in the previous subproblem. 

**Hint:** you may ﬁnd it useful to deﬁne a generator that enumerates the positive integers.

<!-- grade_id: cell-94f79e207c8683e2, points: 0, grade: false, solution: true -->
```python
m=posint()
octa=((i*(i+1)*(2*i+1)+(i-1)*(i)*(2*(i-1)+1))/6 for i in m)
```

Output:
```
1:2: E225 missing whitespace around operator
2:5: E225 missing whitespace around operator
```

<!-- grade_id: cell-c0ec74ed06e7b6e9, points: 1, grade: true, solution: false -->
```python
"""Test your octa generator is correct"""
octa_list = [next(octa) for _ in range(10)]
assert octa_list[:3] == [1, 6, 19]
```

## Problem 5: Map, Filter and Reduce (5 points) 
In this exercise, you’ll learn a bit about map, ﬁlter and reduce operations. We will revisit these operations in a few weeks when we discuss MapReduce and related frameworks in distributed computing. In this problem, I expect that you will use only the functions `map`, `filter` and functions from the `functools` and `itertools` modules, along with the `range` function (and similar list-related functions) and a sprinkling of `lambda` expressions.

**Hint:** You are encouraged to take a look ath the documentation of [functools](https://docs.python.org/3/library/functools.html) and [itertools](https://docs.python.org/2/library/itertools.html) and try to come up with __one-line__ expressions for the following problems based on the functions that you find in the documentation. Useful functions include `functools.reduce` and `itertools.islice`

```python
# you will need these imports
import functools
import itertools
```

### Part 1 (1 point) 
Write a one-line expression that computes the sum of the ﬁrst 10 odd numbers (starting from 1). For ease of grading, please assign the output of this expression to a variable called `sum_of_odd_numbers`.

<!-- grade_id: cell-31bad9a2f9abd832, points: 0, grade: false, solution: true -->
```python
"""Create your variable `sum_of_odd_numbers` here."""
sum_of_odd_numbers=functools.reduce(lambda x,y:x+(2*y+1), range(0,10))
```

<!-- grade_id: cell-85d5fe5a3312e8a6, points: 1, grade: true, solution: false -->
```python
"""Test that your variable is correct"""
assert type(sum_of_odd_numbers) == int
```

### Part 2 (1 point) 
Write a one-line expression that computes the product of the first 13 primes. You may use the `primes` generator that you defined above. Please assign the output of this expression to a variable called `product_of_primes`. 

<!-- grade_id: cell-fee58c92dd3faabc, points: 0, grade: false, solution: true -->
```python
"""Create your variable `product_of_primes` here."""
#use itertools.islice to select elements from the iterable
product_of_primes=functools.reduce(lambda x,y:x*y, itertools.islice(primes(),0,13))
```

<!-- grade_id: cell-39c66a19d8db71ab, points: 1, grade: true, solution: false -->
```python
"""Test your variable is correct"""
assert type(product_of_primes) == int
```

### Part 3 (1 point) 
Write a one-line expression that computes the sum of the squares of the first 31 primes. You may use the `primes` generator that you defined above. Please assign the output of this expression to a variable called `squared_primes`.

<!-- grade_id: cell-73a030197cef30c0, points: 0, grade: false, solution: true -->
```python
"""Create your variable `squared_primes` here."""
squared_primes=functools.reduce(lambda x,y:x+y**2,itertools.islice(primes(),31),0)
```

<!-- grade_id: cell-3a93bee591d5ff41, points: 1, grade: true, solution: false -->
```python
"""Check your variable is correct"""
assert type(squared_primes) == int
```

### Part 4 (1 point) 
Write a one-line expression that computes a list of the ﬁrst twenty harmonic numbers. Recall that the $n$-th harmonic number is given by $H_{n} = \Sigma^{n}_{k=1} 1/k$. For ease of grading, please assign the output of this expression to a variable called `harmonics`.

<!-- grade_id: cell-799635e30e73147f, points: 0, grade: false, solution: true -->
```python
"""Create your variable `harmonics` here"""
harmonics=list(itertools.accumulate(range(1,21),lambda x,y:x+1/y))
```

<!-- grade_id: cell-c2e881166e69bbc7, points: 1, grade: true, solution: false -->
```python
"""Tests for your `harmonics` variable."""
assert harmonics[0:3] == [1.0, 1.5, 1.8333333333333333]
```

### Part 5 (1 point) 
Write a one-line expression that computes the geometric mean of the ﬁrst 12 octahedral numbers. You may use the generator that you wrote in the previous problem. Recall that the _geometric mean_ of a collection of $n$ numbers $a_1 , a_2 , ..., a_{n}$ is given by $(\Pi_{i=1}^{n} a_{i})^{1/n}$. For ease of grading, please assign the output of this expression to a variable called `octa_geom`.

<!-- grade_id: cell-3e31eed02329dcef, points: 0, grade: false, solution: true -->
```python
"""Code to create your variable `octa_geom` goes here."""

octa_geom=functools.reduce(lambda x,y:x*y, [((i*(i+1)*(2*i+1)+(i-1)*(i)*(2*(i-1)+1))/6)**(1/12) for i in range(1,13)])
```

Output:
```
3:10: E225 missing whitespace around operator
3:36: E231 missing whitespace after ','
38: E231 missing whitespace after ':'
3:113: E231 missing whitespace after ','
```

<!-- grade_id: cell-ae19394061803df3, points: 1, grade: true, solution: false -->
```python
"""Tests your `octa_geom`"""
assert type(octa_geom) == float
```

## Problem 6: Fun with Polynomials (3 points)
In this exercise you’ll get a bit of experience writing higher-order functions. You may ignore error checking in this problem.

### Part 1 (2 points) 
Write a function `make_poly` that takes a list of numbers (`int`s and/or `ﬂoat`s) `coeffs` as its only argument and returns a function `p`. The list coeffs encodes the coeﬃcients of a polynomial, $p(x) = a_{0} + a_{1} x + a_{2} x^{2} + ... + a_{n} x_{n}$ , with $a_i$ given by `coeffs[i]`. The function `p` should take a single number (`int` or `ﬂoat`) x as its argument, and return the value of the polynomial `p` evaluated at `x`.

<!-- grade_id: cell-a0c33bf1b4c235e1, points: 0, grade: false, solution: true -->
```python
def make_poly(coeffs):
    def p(x):
        return sum([coeffs[i]*x**i for i in range(len(coeffs))])
    return p
```

<!-- grade_id: cell-1579955824e05d5e, points: 1, grade: true, solution: false -->
```python
"""Tests for make_poly"""
assert make_poly([1, 3, 4])(2) == 23
```

<!-- grade_id: cell-ecb06b3a845aa5ae, points: 1, grade: true, solution: false -->
```python
assert make_poly([1, 2])(3) == 7
```

### Part 2 (1 point) 
Write a function `eval_poly` that takes in a list of coefficients `coeffs` and a list of arguments `args`. Your function should return the result of evaluating the polynomial that has coefficients `coeffs` on each of the values of `args` as a list. 

<!-- grade_id: cell-ee29357cfde3dbca, points: 0, grade: false, solution: true -->
```python
def eval_poly(coeffs, args):
    return list(map(make_poly(coeffs), args))
```

<!-- grade_id: cell-a77bcf64f691cba2, points: 1, grade: true, solution: false -->
```python
"""Tests for eval_poly"""
assert eval_poly([1, 3, 4], [2, 3]) == [23, 46]
```